<a href="https://colab.research.google.com/github/aloesther/DSA2025Test/blob/main/DSA2025Test.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("DSA2025Test.ipynb")
```

# DSA 2025 Summer School Admittance Check

Thanks for your interest in attending DSA 2025 Ibadan, Nigeria. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Ibadan!

## Instructions
1. Complete each function according to the provided specifications
2. Run all cells to verify your solutions
3. All tests must pass to generate a submission
4. Save your work before submitting

```python
# Run these once ... To be on a safe side
!pip install nose
!pip install otter-grader
import IPython
from IPython import get_ipython
# Import the good stuff
import pandas as pd
import numpy as np
import math
from nose.tools import assert_equal
import otter
from collections import Counter
# Assuming your tests are in a directory called "tests"
# Create the directory if it doesn't exist
import os
if not os.path.exists("tests"):
    os.makedirs("tests")

grader = otter.Notebook("DSA2025Test.ipynb", tests_dir= "tests")
```

Output:
```
Requirement already satisfied: nose in /usr/local/lib/python3.11/dist-packages (1.3.7)
Requirement already satisfied: otter-grader in /usr/local/lib/python3.11/dist-packages (6.1.1)
Requirement already satisfied: click<9.0.0,>=8.1.7 in /usr/local/lib/python3.11/dist-packages (from otter-grader) (8.1.8)
Requirement already satisfied: dill>=0.3.0 in /usr/local/lib/python3.11/dist-packages (from otter-grader) (0.3.9)
Requirement already satisfied: fica>=0.4.1 in /usr/local/lib/python3.11/dist-packages (from otter-grader) (0.4.1)
Requirement already satisfied: ipylab<2.0.0,>=1.0.0 in /usr/local/lib/python3.11/dist-packages (from otter-grader) (1.0.0)
Requirement already satisfied: ipython in /usr/local/lib/python3.11/dist-packages (from otter-grader) (7.34.0)
Requirement already satisfied: ipywidgets<9.0.0,>=8.1.5 in /usr/local/lib/python3.11/dist-packages (from otter-grader) (8.1.5)
Requirement already satisfied: jinja2<4.0,>=3.1 in /usr/local/lib/python3.11/dist-packages (from otter-grad
```

<!-- BEGIN QUESTION -->

**Question 1:**
Write a Python function to return a tuple of even and odd numbers

1.   List item
2.   List item

given an integer input.</br>
Example: For input 10, the output should be ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9]).

```python
def even_odd_numbers(n):
    """
    Returns a tuple of even and odd numbers up to n.

    Args:
        n (int): The upper limit.

    Returns:
        tuple: A tuple of two lists: (evens, odds).
    """
    evens = [i for i in range(1, n + 1) if i % 2 == 0]
    odds = [i for i in range(1, n + 1) if i % 2 != 0]

    """
    Returns a tuple containing a list of even numbers and a list of odd numbers up to n (inclusive).
    """
    return (evens, odds)
# Example usage:
print(even_odd_numbers(10))


```

Output:
```
([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2: Create a dictionary to store the frequency of each word in the paragraph provided.**</br>
Example: For the paragraph "Data science is fun. Data science is useful.", the output should be {'Data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}.

```python
grader.check("q2")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3: Extract the values from the dictionary above into a list and sort them in descending order.**</br>
Example: For the dictionary {'a': 3, 'b': 1, 'c': 2}, the output should be [3, 2, 1]

```python
def extract_and_sort_values(dictionary):
    """
    Extracts and sorts dictionary values in descending order.

    Args:
        dictionary (dict): The input dictionary.

    Returns:
        list: A list of sorted values in descending order.
    """
    value_list = [num for num in dictionary.values()]
    sorted_list = sorted(value_list, reverse=True)
    return sorted_list

    # Example dictionary
example_dict = {'a': 3, 'b': 1, 'c': 2}

# Expected output: [3, 2, 1]
print(extract_and_sort_values(example_dict))


```

Output:
```
[3, 2, 1]
```

```python
grader.check("q3")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4: Merge the two lists (from Questions 3 and 4) into a dictionary where the keys are the sorted keys and the values are the sorted values.**</br>
Example: For lists ['a', 'b', 'c'] and [3, 2, 1], the output should be {'a': 3, 'b': 2, 'c': 1}

```python
def merge_lists_to_dict(dictionary):  # Changed to accept a dictionary
    """
    Merges dictionary keys and sorted values into a new dictionary.

    Args:
        dictionary (dict): The input dictionary.

    Returns:
        dict: A dictionary with sorted keys and values paired.
    """
    keys = sorted(dictionary.keys())
    values = sorted(dictionary.values(), reverse=True)
    return dict(zip(keys, values))

# Example usage:
print(merge_lists_to_dict(example_dict))  # Call with the dictionary
```

Output:
```
{'a': 3, 'b': 2, 'c': 1}
```

```python
grader.check("q4")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5:Write a function least_common_multiple that takes two inputs a and b and returns the least common multiple of the two numbers.**</br>
Example: For input (4, 6), the output should be 12

```python
def least_common_multiple(a, b):
    """
    Returns the least common multiple of two numbers.

    Args:
        a (int): First number.
        b (int): Second number.

    Returns:
        int: The LCM of a and b.
    """
    multiple = a * b
    # Use the greatest common divisor function from the maths module
    greatest_common_divisor = math.gcd(a,b)

    # Using floor division in order to get a integer output
    lcm = multiple // greatest_common_divisor
    return lcm


```

```python
grader.check("q5")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6:Write a function get_nearest_farthest that takes in a point of interest (x, y) and a list of points [(x1, y1), (x2, y2), ...] and returns the indices of the nearest and farthest points from the point of interest.** </br>
Example: For pt = (0, 0) and points = [(1, 1), (3, 3), (-1, -1)], the output should be (0, 1)

```python
def get_nearest_farthest(pt, points):

    """
    Returns the indices of the nearest and farthest points from a point of interest.

    Args:
        pt (tuple): The point of interest (x, y).
        points (list): List of points [(x1, y1), (x2, y2), ...].

    Returns:
        tuple: Indices of the nearest and farthest points.
    """
    distances = []
    for point in points:
        dist = abs(np.array(pt) - np.array(point))
        dist_magnitude = np.sum(dist)
        distances.append(dist_magnitude)

    nearest_point = min(distances)
    farthest_point = max(distances)

    return (distances.index(nearest_point), distances.index(farthest_point))

```

```python
grader.check("q6")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:Write a function filter_divisible to return a list of numbers between 0 and a number N that are not perfectly divisible by q.** </br>

> Hint: If N is negative, use N = 20.

Example: For N = 10 and q = 2, the output should be [1, 3, 5, 7, 9]

```python
def filter_divisible(N, q):
    """
    Returns a list of numbers between 0 and N that are not divisible by q.

    Args:
        N (int): The upper limit.
        q (int): The divisor.

    Returns:
        list: A list of non-divisible numbers.
    """
    non_divisible_nums = []
    for i in range(N+1):
        if i % q != 0:
            non_divisible_nums.append(i)

    return non_divisible_nums
```

```python
grader.check("q7")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8: Write a function flatten_and_unique that takes in a list of lists and outputs a sorted list of unique elements from all sublists.**

```python
def flatten_and_unique(list_of_lists):
    """
    Flattens a list of lists and returns a sorted list of unique elements.

    Args:
        list_of_lists (list): A list of lists.

    Returns:
        list: A sorted list of unique elements.
    """
    all_elements = []
    for sublist in list_of_lists:
        for element in sublist:
            all_elements.append(element)
    unique_elements = set(all_elements)
    unique_list = list(unique_elements)
    return sorted(unique_list)
```

```python
grader.check("q8")
```

<!-- END QUESTION -->

**The Extra Mile!**


Download the dataset Nigeria Food Prices (9.9M)

> https://data.humdata.org/dataset/wfp-food-prices-for-nigeria





<!-- BEGIN QUESTION -->

**Question 9: Create a new column date_new from the date column, converting it to a datetime format.**

```python
import pandas as pd

# Load the dataset
file_path = '/content/wfp_food_prices_nga.csv'  # Replace with the actual file path
df = pd.read_csv(file_path,skiprows=[1])
df = df.drop(index=2)
df = df.reset_index(drop=True)

# Display the first few rows of the dataset
print("First 5 rows of the dataset:")
print(df.head())


```

Output:
```
First 5 rows of the dataset:
         date   admin1 admin2        market  latitude  longitude  \
0  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
1  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
2  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
3  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
4  2002-01-15   Sokoto   Gada  Illela (CBM)    13.645      5.278   

             category      commodity unit priceflag  pricetype currency  \
0  cereals and tubers          Maize   KG    actual  Wholesale      NGN   
1  cereals and tubers         Millet   KG    actual  Wholesale      NGN   
2  cereals and tubers        Sorghum   KG    actual  Wholesale      NGN   
3     pulses and nuts  Beans (niebe)   KG    actual  Wholesale      NGN   
4  cereals and tubers          Maize   KG    actual  Wholesale      NGN   

    price  usdprice  
0  175.92    1.5398  
1  150.18    1.3145  
2  155.61    1.3620  
3  196.87    1.7231  
4
```

```python
def create_date_new(df):
    """
    Creates a new column `date_new` from the `date` column, converting it to a datetime format.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.DataFrame: The dataframe with the new `date_new` column.
    """
    df['date_new'] = pd.to_datetime(df["date"])

    return df
```

```python
grader.check("q9")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10:Split the dataframe into two separate dataframes based on the category column: one for cereals and tubers and another for pulses and nuts**

```python
def split_dataframes(df):
    """
    Splits the dataframe into two separate dataframes based on the `category` column.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        tuple: A tuple of two dataframes: (cereals_df, pulses_df).
    """
    cereals_df = df[df.category == "cereals and tubers"]
    pulses_df = df[df.category == "pulses and nuts"]

    return (cereals_df, pulses_df)
```

```python
grader.check("q10")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11:Calculate the mean, median, and mode of the price and usdprice columns for each category**

```python
def calculate_price_stats(df):
    """
    Calculates the mean, median, and mode of the `price` and `usdprice` columns for each `category`.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        tuple: A tuple of two dataframes: (price_stats, usdprice_stats).

    """

    price_stats = pd.DataFrame(data = {"mean" : df["price"].mean(), "median":df["price"].median(), "mode":df["price"].mode()})
    usdprice_stats = pd.DataFrame(data = {"mean" : df["usdprice"].mean(), "median":df["usdprice"].median(), "mode":df["usdprice"].mode()})
    return (price_stats, usdprice_stats)
```

```python
grader.check("q11")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12:Merge the two dataframes (from Question 11) back into one dataframe and save it as `merged_data.csv`.**

```python
def merge_dataframes(cereals_df, pulses_df):
    """
    Merges the two dataframes back into one dataframe and saves it as `merged_data.csv`.

    Args:
        cereals_df (pd.DataFrame): The cereals and tubers dataframe.
        pulses_df (pd.DataFrame): The pulses and nuts dataframe.

    Returns:
        pd.DataFrame: The merged dataframe.
    """
    merged_df = pd.concat([cereals_df, pulses_df])
    return merged_df.reset_index(drop=True)
```

```python
grader.check("q12")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13:Open the `merged_data.csv` file and select only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns**

```python
def select_columns():
    """
    Opens the `merged_data.csv` file and selects only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns.

    Returns:
        pd.DataFrame: The dataframe with selected columns.
    """
    data = pd.read_csv('merged_data.csv', usecols=["date_new", "market", "commodity", "price", "usdprice"])

    return data
```

```python
grader.check("q13")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14:Group the data by `admin1` (state) and calculate the average `price` and `usdprice` for each state**

```python
def calculate_state_avg_prices(df):
    """
    Groups the data by `admin1` (state) and calculates the average `price` and `usdprice` for each state.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.DataFrame: The dataframe with average prices by state.
    """
    grouped_by_state = df.groupby("admin1")[["price", "usdprice"]].mean()

    return grouped_by_state
```

```python
grader.check("q14")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 15:Identify the top 5 markets with the highest average `usdprice` for `Rice (imported)`.**

```python
def identify_top_markets(df):
    """
    Identifies the top 5 markets with the highest average `usdprice` for `Rice (imported)`.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.Series: A series with the top 5 markets and their average `usdprice`.
    """
    rice = df[df.commodity == "Rice (imported)"]
    grouped_df = rice.groupby(by="market")["usdprice"].mean().sort_values(ascending=False)
    top_5_markets = grouped_df.head(5)
    return top_5_markets
```

```python
grader.check("q15")
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Complete each function according to the provided specifications. Make sure all tests pass.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
/usr/local/lib/python3.11/dist-packages/otter/check/utils.py in grading_mode_disabled(wrapped, self, args, kwargs)
    204     if type(self)._grading_mode:
    205         return
--> 206     return wrapped(*args, **kwargs)
    207 
    208 

/usr/local/lib/python3.11/dist-packages/otter/check/utils.py in incompatible(wrapped, self, args, kwargs)
    181         if self.interpreter is interpreter:
    182             if throw_error:
--> 183                 raise RuntimeError(
    184                     f"This method is not compatible with {interpreter.value.display_name}"
    185                 )

RuntimeError: This method is not compatible with Google Colab
```

Tests for q1 (points: 1)
```
>>> assert even_odd_numbers(10) == ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```
```
>>> assert even_odd_numbers(5) == ([2, 4], [1, 3, 5])
```


Tests for q10 (points: 1)
```
>>> (cereals_df, pulses_df) = split_dataframes(df)
>>> assert len(cereals_df) > 0, 'Cereals and tubers dataframe is empty'
```
```
>>> assert len(pulses_df) > 0, 'Pulses and nuts dataframe is empty'
```


Tests for q11 (points: 1)
```
>>> (price_stats, usdprice_stats) = calculate_price_stats(df)
>>> assert not price_stats.empty, 'Price statistics dataframe is empty'
```
```
>>> assert not usdprice_stats.empty, 'USD price statistics dataframe is empty'
```


Tests for q12 (points: 1)
```
>>> merged_df = merge_dataframes(cereals_df, pulses_df)
>>> assert not merged_df.empty, 'Merged dataframe is empty'
```
```
>>> assert len(merged_df) == len(cereals_df) + len(pulses_df), 'Merged dataframe has incorrect number of rows'
```


Tests for q13 (points: 1)
```
>>> selected_columns = select_columns()
>>> assert not selected_columns.empty, 'Selected columns dataframe is empty'
```
```
>>> assert list(selected_columns.columns) == ['date_new', 'market', 'commodity', 'price', 'usdprice'], 'Incorrect columns selected'
```


Tests for q14 (points: 1)
```
>>> state_avg_prices = calculate_state_avg_prices(df)
>>> assert not state_avg_prices.empty, 'State average prices dataframe is empty'
```
```
>>> assert 'price' in state_avg_prices.columns, 'Price column not found in state average prices'
```


Tests for q15 (points: 1)
```
>>> top_markets = identify_top_markets(df)
>>> assert not top_markets.empty, 'Top markets dataframe is empty'
```
```
>>> assert len(top_markets) == 5, 'Incorrect number of top markets'
```


Tests for q2 (points: 1)
```
>>> assert word_frequency('Data science is fun. Data science is useful.') == {'data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}
```
```
>>> assert word_frequency('Hello hello world') == {'hello': 2, 'world': 1}
```


Tests for q3 (points: 1)
```
>>> assert extract_and_sort_values({'a': 3, 'b': 1, 'c': 2}) == [3, 2, 1]
```
```
>>> assert extract_and_sort_values({'x': 10, 'y': 5, 'z': 20}) == [20, 10, 5]
```


Tests for q4 (points: 1)
```
>>> assert merge_lists_to_dict(['a', 'b', 'c'], [3, 2, 1]) == {'a': 3, 'b': 2, 'c': 1}
```
```
>>> assert merge_lists_to_dict(['x', 'y'], [10, 20]) == {'x': 10, 'y': 20}
```


Tests for q5 (points: 1)
```
>>> assert least_common_multiple(4, 6) == 12
```
```
>>> assert least_common_multiple(5, 10) == 10
```


Tests for q6 (points: 1)
```
>>> assert get_nearest_farthest((0, 0), [(1, 1), (3, 3), (-1, -1)]) == (0, 1)
```
```
>>> assert get_nearest_farthest((0, 0), [(2, 2), (1, 1), (3, 3)]) == (1, 2)
```


Tests for q7 (points: 1)
```
>>> assert filter_divisible(10, 2) == [1, 3, 5, 7, 9]
```
```
>>> assert filter_divisible(15, 3) == [1, 2, 4, 5, 7, 8, 10, 11, 13, 14]
```


Tests for q8 (points: 1)
```
>>> assert flatten_and_unique([[2, 13, 44], [6, 7, 2]]) == [2, 6, 7, 13, 44]
```
```
>>> assert flatten_and_unique([[1, 2, 3], [3, 4, 5]]) == [1, 2, 3, 4, 5]
```


Tests for q9 (points: 1)
```
>>> df = create_date_new(df)
>>> assert 'date_new' in df.columns, "Column 'date_new' not found in the dataframe"
```
```
>>> assert df['date_new'].dtype == 'datetime64[ns]', "Column 'date_new' is not in datetime format"
```
